```python
# You must run this cell, but you can ignore its contents.

import hashlib

def ads_hash(ty):
    """Return a unique string for input"""
    ty_str = str(ty).encode()
    m = hashlib.sha256()
    m.update(ty_str)
    return m.hexdigest()[:10]
```

# More fun with function arguments

## Treating arguments differently depending on the number given

We have dealt with several functions which work differently depending on how many arguments they are called with. For example, the Python built-in function `range()` treats its arguments differently, depending on how many are given.

- `range(10)` returns an iterator over integers, starting from 0 and stopping before 10.

- `range(2, 10)` returns an iterator over integers, starting from 2 and stopping before 10.

So, the function handles the first argument as the stop index if only one argument is given. In this case, the start index defaults to a value of 0.

If two arguments are given, the first argument is the start index and the second argument is the stop index.

We could make our own function that behaves similarly.


```python
def range_list(arg0, arg1=None):
    if arg1 == None:
        # only one argument given
        start_idx = 0        
        stop_idx = arg0
    else:
        start_idx = arg0
        stop_idx = arg1
    result = []
    i = start_idx
    while i < stop_idx:
        result.append(i)
        i = i + 1
    return result
```

```python
range_list(10)
```

```python
range_list(2,10)
```

## Q1 Treating arguments differently depending on the number of arguments given

Create a function called `get_2D_coordinates` which returns a list of `(x, y)` tuples.

If called with a single argument, this function will treat this as a sequence of `y` values and return a list of tuples where `x` starts at `0` and continues to `len(arg0)-1`.

If called with two arguments, this function will treat the first argument as a sequence of `x` values and the second argument as a sequence of `y` values. Both input arguments must be the same length in this case.

For example, if you call your function like this: `get_2D_coordinates([1,2,3])`, it should return `[(0, 1), (1, 2), (2, 3)]`. If you call it like this `get_2D_coordinates([1,2,3],[4,5,6])`, it should return `[(1, 4), (2, 5), (3, 6)]`.

<!-- grade_id: cell-04cf7976bab204cf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-025802508de10ec1, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(get_2D_coordinates([10,20,30,40]))=='669a1e9cc7'
assert ads_hash(get_2D_coordinates([2,5,9]))=='f2aeea922b'
```

# numpy

```python
import numpy as np
```

## Q2 Create a 1 dimensional numpy array named `x` with 20 elements from 0 to 19.

<!-- grade_id: cell-cff3a5e50b06014a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-64723536c462a2d1, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(x)=='7a150607a7'
```

## Q3 Create a 2 dimensional numpy array named `x` of shape 5,6.

<!-- grade_id: cell-15f8e74190eecaf8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-38a672b790c362eb, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(x.shape)=='6510380315'
```

## Q4 Consider the first 100 integers starting with 0. What is their mean value? Put this in a variable `mean100`.

<!-- grade_id: cell-34f66fe8d8b001fb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9ef246cf9714c370, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(round(mean100*1000))=='bc8d4a4d1e'
```

## Q5 Create an array named `x` with every 4th value of the first 100000 integers.

<!-- grade_id: cell-6b4ee0992aa506c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6833bfed10c7430c, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(len(x))=='0812a4ef4e' 
assert ads_hash(x)=='17e101b2ef'
```

## Q6 set the shape of `x` to be (5, 5000)

<!-- grade_id: cell-e07f245d339e6328, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-190e11601c41f920, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(len(x))=='ef2d127de3' 
assert ads_hash(x)=='3331e0302b'
```

## Q7 Create the variable `y` from the first 3 rows and the first 8 columns of `x`.

<!-- grade_id: cell-ef33bd87060e2253, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-38599e5840dfe50f, points: 0, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(len(y))=='4e07408562' 
assert ads_hash(y)=='b39d640ca0'
```

## Plotting numpy arrays with matplotlib

Let's start plotting numpy arrays. The `plot()` function in matplotlib's `pyplot` module works differently depending on the number of arguments it was called with. Like your `get_2D_coordinates()` function above, if only one argument is used, it is treated as the `y` coordinates. The `x` coordinates are the integers starting with `0`. If two arguments are used, they are treated as the `x` and `y` coordinates.

Here we are also going to label our axes to keep track of what we are plotting. We will use the `plt.xlabel()` and `plt.ylabel()` functions.

```python
import matplotlib.pyplot as plt
```

We are going to generate some "toy" (artificial, false, fake, made-up) positions of a bee over time. First, we need the times where we have tracking data. We are going to track the bee for 10 seconds. Over these 10 seconds, we have 100 measurements.

```python
time_values = np.linspace(0,10,100)
plt.plot(time_values)
plt.ylabel('time (seconds)');
```

Now we will get the `y` coordinate of the bee over these 10 seconds. Instead of real data, we generate toy data by using the `np.sin()` function. We take the sine of the time value as the `y` coordinate of the bee.

```python
y = np.sin(time_values)
```

Now let's plot the `y` coordinate of the bee over time.

```python
plt.plot(time_values,y)
plt.xlabel('time (seconds)');
plt.ylabel('y (meters)');
```

So our bee was oscillating in `y` over time.

In addition to a `y` coordinate, our tracking system also gives us `x` coordinates. Let's make some toy data for this,  too.

```python
x = 0.5*np.sin(time_values*3.2)+1.1
```

```python
plt.plot(time_values,x)
plt.xlabel('time (seconds)');
plt.ylabel('x (meters)');
```

So the bee was also oscillating in `x`, and even faster than in `y`.

## Storing multiple dimensions data in single array

We might want to store our coordinates in a single multi-dimensional numpy array where the columns are time and the rows are the different coordinates (`x`, and `y`).

```
x0, x1, x2, ..., xn
y0, y1, y2, ..., yn
```

We can use the `np.vstack()` ("vertical stack") function for this. It will return a new array whose rows are the given in its input argument.

```python
bee_coords = np.vstack((x,y))
```

By doing these, we can keep the different measurements together. Of course, we can still plot just a single dimension. Here is `x` over time again. In this case, the x axis is simply the index of the array passed to `plot()`:

```python
plt.plot(bee_coords[0])
plt.ylabel('x (meters)');
```

Of course, we can also plot the time value as the x axis and label it correctly:

```python
plt.plot(time_values, bee_coords[0])
plt.xlabel('time (seconds)');
plt.ylabel('x (meters)');
```

If `x` and `y` are horizontal positions, we can plot a "top view" of the bee position by plotting x against y:

```python
plt.plot(bee_coords[0], bee_coords[1])
plt.xlabel('x (meters)')
plt.ylabel('y (meters)');
```

Let's now calculate how far the bee is in any moment in time from a particular location. Remember the distance function we made in lecture:

```python
def compute_dist(a,b):
    """compute distance between two points, `a` and `b`
    
    `a` and `b` are each 2D sequences representing the coordinates of the point."""
    dx = b[0] - a[0]
    dy = b[1] - a[1]
    return np.sqrt( dx**2 + dy**2 )
```

We wrote this function with the idea that `a` and `b` are each a 2D point. With `a = (xa, ya)` and `b = (xb, yb)`.

Now we want a a function which returns the distance between a single input point `a` with coordinates `(xa, ya)` and each and every point in an array `bee_coords` which has shape `(2, n)`. So column `i` of `bee_coords` is thus a coordinate `(xi, yi)`.

Let's build up to this function piece-by-piece.

## Q8

First write a function called `computed_dx_squared` with the signature:

```python
def computed_dx_squared(a,b):
    """return the x distance squared between `a` and all points in `b`.
    
    `a` is a length 2 sequence with coordinates `(xa, ya)` and `b` is a numpy array of shape `(2, n)`.
    """
```

This function returns the squared `x` distance between point `a` (a length 2 sequence of the `x`,`y` coordinates of point `a`) and each coordinate in the 2D array `b` with shape `(2,n)`. Thus, for `a=[1.0,2.1]` and `b=np.array([[1,2,3],[4,5,6]])`, this will return `np.array([0., 1., 4.])`.

<!-- grade_id: cell-68e7d25afb2f1765, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-917fe5138adc3d85, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(computed_dx_squared([1.0,2.1],np.array([[1,2,3],[4,5,6]])))=='1bacf727ce'
assert ads_hash(computed_dx_squared([1.0,2.1],np.array([[1,2,3,4.2],[4,5,6,8.2]])))=='c9c0596224'
```

## Q9

Now write a function called `computed_dy_squared` with the signature:

```python
def computed_dy_squared(a,b):
    """return the y distance squared between `a` and all points in `b`.
    
    `a` is a length 2 sequence with coordinates `(xa, ya)` and `b` is a numpy array of shape `(2, n)`.
    """
```

This function returns the squared `y` distance between point `a` (a length 2 sequence of the `x`,`y` coordinates of point `a`) and each coordinate in the 2D array `b` with shape `(2,n)`. Thus, for `a=[1.0,2.1]` and `b=np.array([[1,2,3],[4,5,6]])`, this will return `array([ 3.61,  8.41, 15.21])`.

<!-- grade_id: cell-e74d9319eda293c6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-52b68c1bd8f83a48, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(computed_dy_squared([1.0, 2.1], np.array([[1,2,3],[4,5,6]])))=='759d96c1e4'
assert ads_hash(computed_dy_squared([1.0, 2.1], np.array([[1,2,3,1.2],[4,5,6,-3]])))=='cb574f3c8e'
```

## Q10

Now write a function called `computed_distance` with the signature:

```python
def compute_distance(a,b):
    """return the distance squared between `a` and all points in `b`.
    
    `a` is a length 2 sequence with coordinates `(xa, ya)` and `b` is a numpy array of shape `(2, n)`.
    """
```

This function returns the distance between point `a` (a length 2 sequence of the `x`,`y` coordinates of point `a`) and each coordinate in the 2D array `b` with shape `(2,n)`. Thus, for `a=[1.0,2.1]` and `b=np.array([[1,2,3],[4,5,6]])`, this will return `array([1.9       , 3.06757233, 4.3829214 ])`.

<!-- grade_id: cell-cbbcbdf110a030b2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eb44de204d241fa6, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(compute_distance([1.0, 2.1], np.array([[1,2,3],[4,5,6]])))=='1e2fc6c5fd'
assert ads_hash(compute_distance([1.0, 2.1], np.array([[1,2,3,10.1],[4,5,6,-2.2]])))=='b7ef5a80ae'
```

## Q11

Now let's imagine you upgraded your tracking system to work in 3 dimensions and now you collect data in `x`, `y` and `z`. Write a function called `compute_distance_3d` with the signature:

```python
def compute_distance_3d(a,b):
    """return the distance squared between `a` and all points in `b`.
    
    `a` is a length 3 sequence with coordinates `(xa, ya, za)` and `b` is a numpy array of shape `(3, n)`.
    """
```

This function returns the distance between point `a` (a length 3 sequence of the `x`,`y`,`z` coordinates of point `a`) and each coordinate in the 2D array `b` with shape `(3,n)`. Thus, for `a=[1.0,2.1,0.0]` and `b=np.array([[1,2,3],[4,5,6],[7,8,9]])`, this will return `np.array([ 7.25327512,  8.56796359, 10.01049449])`.

Remember that the formula for distance in 3D is:

```
distance = sqrt( dx**2 + dy**2 + dz**2 )
```

where `dx` is the distance along the `x` direction, `dy` is the distance along the `y` direction, and `dz` is the distance along the `z` direction, 

<!-- grade_id: cell-d3cd40a661182621, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f0538c2e9972a4ba, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(compute_distance_3d([1.0,2.1,0.0], np.array([[1,2,3],[4,5,6],[7,8,9]])))=='6f84b17de0'
assert ads_hash(compute_distance_3d([1.0,2.1,0.1], np.array([[1,2,3,0.0],[4,5,6,1.0],[7,8,9.2,0.0]])))=='0b9c9c17b1'
```