```python
# You must run this cell, but you can ignore its contents.

import hashlib

def ads_hash(ty):
    """Return a unique string for input"""
    ty_str = str(ty).encode()
    m = hashlib.sha256()
    m.update(ty_str)
    return m.hexdigest()[:10]
```

```python
# You must also run this cell.

import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
```

# Optimization first steps

We are going to take our first steps towards optimization by returning to a bumblebee example.

We are going to define the positions of a flower and the flight path of a bumblebee.

```python
flower = np.array([7.5, 10.3])
```

```python
def make_bee_track(t):
    pos0 = (-10,3)
    velocity = (4.0, 0.2)
    pos_x = pos0[0] + t*velocity[0]
    pos_y = pos0[1] + t*velocity[1]
    return np.array([pos_x,pos_y])

t = np.linspace(0,15,20)
bee_track = make_bee_track(t)
```

Here we plot these positions.

```python
fig, ax = plt.subplots(nrows=1, ncols=1)
ax.plot( [flower[0]], [flower[1]], 'or', label='flower' )
ax.plot( bee_track[0], bee_track[1], '.-k', label='bee')
ax.axis('equal')
ax.legend();
```

## Q1 Draw a figure as above with, additionally, a blue line between each point on `bee_track` and `flower`.

When complete, your figure should look like this:

![bee-dist.png](bee-dist.png)

<!-- grade_id: cell-43f8774b1a82b15b, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Q2 compute the distance between each point on `bee_track` and `flower`. Put the results in a 1D numpy array called `distance`.

Hint: recall the function you wrote in the "bumblebee traplines" exercise called `my_distance`.

<!-- grade_id: cell-bf78a3429e48b28d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-70fd1e669decb646, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(np.round(distance*1000).astype(np.int))=='54f4f2edcb'
```

## Q3 make a plot of the bee track parameter `t` on the X axis and `distance` on the Y axis.

It should look like this:

![t-vs-dist.png](t-vs-dist.png)

<!-- grade_id: cell-acdc61a7b1539944, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Q4 Using `DistanceFinder` from the lecture, find the value of `t` that minimizes the distance between the bee and the flower. Save the result in `best_t`.

<!-- grade_id: cell-23d19427610d2072, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0a3594afd2952a33, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(np.round(best_t*1000).astype(np.int))=='dec1ab2f6d'
```

```python
make_bee_track(best_t)
```

## Q5 What is the position of the bee when it is closest to the flower? Save the result as a numpy array in `best_pos`.

<!-- grade_id: cell-cac598c7f739d146, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9d882991838f0996, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert type(best_pos)==np.ndarray
assert best_pos.ndim==1
assert best_pos.shape==(2,)
assert ads_hash(np.round(best_pos[0]*1000).astype(np.int))=='e33b9415bc'
assert ads_hash(np.round(best_pos[1]*1000).astype(np.int))=='f71cbfce4c'
```