Login to your account at https://strawlab-rp2.zoologie.uni-freiburg.de/.

Download this notebook from github:

- Option 1) Navigate to view the file in the browser, click the "Raw" view near the upper right, then in your browser save the file with "Save" or "Save As" and download the `.ipynb` file.

- Option 2) On the front page of the github repository, click on the green `↓ Code` button near the upper right and pick one of the methods. 'Zip' is easiest to start.

Upload this notebook into Jupyter at https://strawlab-rp2.zoologie.uni-freiburg.de/.

Type "1+1" (without the quotes in the cell below) and run the cell.

<!-- grade_id: cell-89f9c5bdfc7d0a6a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9c2fbb8378f13a03, points: 1, grade: true, solution: false -->
```python
assert(_==2)
```